---
title: sdgsgsdgsd
description: >-
  undefinedsdfgdg


---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 100

key: 902b6d4a70
```

2

`@instructions`
2

`@hint`
2

`@pre_exercise_code`
```{undefined}
2
```
`@sample_code`
```{undefined}
2
```
`@solution`
```{undefined}
2
```
`@sct`
```{undefined}
2
```





---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 100

key: 6175bf5139
```

1

`@instructions`
1

`@hint`
1

`@pre_exercise_code`
```{undefined}
1
```
`@sample_code`
```{undefined}
1
```
`@solution`
```{undefined}
1
```
`@sct`
```{undefined}
1
```



